
# Module 2 Assignment


A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
from nose.tools import assert_equal, assert_true

import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Create variables

In this problem, create 2 varables.
Name the first variable **weight** and set it to 160.
Name the second variable **height** and set it to 72.


<!-- grade_id: p1_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
weight = 160
height = 72
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
assert_equal(weight, 160, msg="weight is not 160")
assert_equal(height, 72, msg="height is not 72")
```

# Problem 2: Calculate BMI

Calculate body mass index(BMI) based on weight and height defined in problem1, then assign the value to variable **bmi**.

The fomula to calculate BMI(assume weight in pounds and height in inches) is: 

$BMI = \large{\frac{weight \times 703}{height^2}}$


<!-- grade_id: p2_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
bmi = (weight*703)/(height*height)
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
assert_equal(round(bmi,1), 21.7, msg="bmi is not correctly calculated" )
```

# Problem 3: Complete function to calculate BMI

- Complete function get_bmi() that accepts two parameters, **weight** and **height**.
- Calculate BMI based on **weight** and **height**(assume weight in pounds and height in inches) in the function
- Return BMI value from the function.

<!-- grade_id: p3_ans, points: 0, grade: false, solution: true -->
```python
def get_bmi(weight, height):
    # YOUR CODE HERE
    return (weight*703)/(height*height)
```

<!-- grade_id: p3_test, points: 1, grade: true, solution: false -->
```python
assert_equal(round(get_bmi(160, 72),1), 21.7, msg="get_bmi() doesn't calculate bmi correctly")
assert_equal(round(get_bmi(180, 68),1), 27.4, msg="get_bmi() doesn't calculate bmi correctly")
```

# Problem 4: Create a function to determine if a number is even
- Write a function called `is_even_number` that accepts one parameters called `num`.   
   
- Inside of this function, return True if num is even number, otherwise return False.

- Hint: Use remainder(%) operator


<!-- grade_id: p4_ans, points: 0, grade: false, solution: true -->
```python
def is_even_number(num):
    '''
    Returns True if num is even number, otherwise False
    
    Parameters
    ----------
    num : an integer
    
    Returns
    -------
    True if num is even number, otherwise False
    '''
    # YOUR CODE HERE
    if num%2 == 0:
        return True
    else:
        return False

```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
assert_true(is_even_number(2), msg='2 is an even number.')
assert_true(not is_even_number(1), msg='1 is not an even number')
```

```python

```