# Module 3 Assignment


A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
import numpy as np
from nose.tools import assert_equal, assert_is_instance

import warnings
warnings.filterwarnings('ignore')
```

# Problem 1: Making a list, tuple, and dictionary

- Write a function called "create_data" which takes in 3 parameters and returns a list, tuple, and dictionary (in that order) containing those 3 arguments.


- The arguments will be labeled **a, b**, and __c__. These arguments should be inputted **in that order** into each data structure (the list, tuple, and dictionary).


- For the dictionary
 - set the key for __a__ to be string "1"(not number 1), 
 - the key for **b** to be string "2", and 
 - the key for __c__ to be string "3".

- Assign the list to variable **my_list**, tuple to variable **my_tuple**, dictionary to variable **my_dict**.

- Make sure that the function returns the list, tuple, and dictionary **in that order**.

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
def create_data(a,b,c):
    '''
    A function which creates a list, tuple, and dictionary of the inputted data.
    
    
    Parameters
    -------
    a: the first datapoint to be inputted
    
    b: the second datapoint to be inputted
    
    c: the third datapoint to be inputted
    
    
    Returns
    --------
    
    my_list: a list containing the datapoints
    
    my_tuple: a tuple containing the datapoints
    
    my_dict: a dictionary containing the datapoints
    
    '''
    
    
    # YOUR CODE HERE
    
    
    return my_list, my_tuple, my_dict
```

```python
a,b,c = 3,4,5
my_list, my_tuple, my_dict = create_data(a,b,c)
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_is_instance(my_list,list,msg='my_list is not a list!')
assert_is_instance(my_tuple,tuple,msg='my_tuple is not a tuple!')
assert_is_instance(my_dict,dict,msg='my_dict is not a dict')

assert_equal(a,my_list[0],msg='my_list has the wrong data, check the order.')
assert_equal(b,my_tuple[1],msg = 'my_tuple has the wrong data, check the order.')
assert_equal(c,my_dict['3'],msg='my_dict either has the wrong data or the keys are not set properly')
```

# Problem 2: Getting elements from a list


- Write a function called "first_thru_20" that takes in a list and returns the first 20 elements of that list, using Python list slicing. Assign the sliced list to variable **first_20** and return it from the function.


<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
def first_thru_20(input_list):
    '''
    A function which takes in a list and returns the first 20 elements
    
    
    Parameters
    -------
    input_list: a Python list
    
    
    Returns
    --------
    
    first_20: a list containing the first 20 elements of the list
    
    '''
    
    
    # YOUR CODE HERE
    
    
    return first_20
```

```python
input_list =list(np.arange(34))
first_20 = first_thru_20(input_list)
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(first_20), list, msg ='Your function did not return a list')
assert_equal(first_20[0]**2, 0, msg ='The first element is incorrect')
assert_equal(first_20[-1]**2, 361, msg = 'The last element is incorrect')
```

# Problem 3: Find the sum from 1 to n

- Write a function called "summer" that takes in a number and uses a Python **for** loop to find the sum of all integers from 1 to that number


- As an example, if the input is 5, then the function should return 1+2+3+4+5 = 15.

- Assign the sum to variable **final_sum** and return it from the function.

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
def summer(n):
    '''
    A function which takes in a number and sums all integers from 1 to that number
    
    
    Parameters
    -------
    n: an integer
        
    
    Returns
    --------
    
    final_sum: the sum of the two lists
    
    '''
    
    
    # YOUR CODE HERE
    
    
    return final_sum
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
temp_sum = summer(2344)
assert_equal(type(temp_sum), int, msg ='Your function did not return an integer')
assert_equal(temp_sum, 2748340, msg ='The sum is incorrect')
```

# Problem 4: Check if a value is in a list

- Write a function called "checker" that takes in a Python list and an integer value, and checks whether or not that value is in the list.


<!-- grade_id: p4_answer, points: 0, grade: false, solution: true -->
```python
def checker(input_list, value):
    '''
    A function which takes in a list and and a value and checks if the value is in the list
    
    
    Parameters
    -------
    input_list: a Python list
    
    value: a value
    
    Returns
    --------
    
    boolean, True if value is in the input list and False otherwise
    
    '''
    
    
    # YOUR CODE HERE
    
    
```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
assert_equal(checker([1,10],9), False,msg='Your function returned an incorrect value')
assert_equal(checker([1,10],10), True,msg='Your function returned an incorrect value')
assert_is_instance(checker([1,10],9), bool,msg='Your function did not return a boolean')
```